# Setup

CS1302 Introduction to Computer Programming
___

## JupyterHub

### How to access the JupyterHub server?

- Enter the url of the Jupyterhub server [divedeep.cs.cityu.edu.hk](https://divedeep.cs.cityu.edu.hk) in a web browser.
- Enter your [EID](https://www.cityu.edu.hk/esu/eid.htm) and Password in the fields `Username` and `Password` respectively.
- Click the `Sign in` button.  

![log in](images/login.dio.svg)

- Choose `CS1302` as the server option and click `Start`.  

![server options](images/serveroptions.dio.svg)

---

**Note**

`Minimal` is only for trouble-shooting and is insufficient for running the course materials. If you accidentally choose it, you can restart your server after entering the JupyterLab interface as described below.

---

- The JupyterLab interface should appear.  

![jupyterlab](images/jupyterlab.dio.svg) 

---

**Note**

To submit homework assignments, you will need to use a different interface:

- `Help`$\to$`Launch Classic Notebook` will open the classic notebook interface, which is needed for submitting lab assignments. You may also
  - visit https://divedeep.cs.cityu.edu.hk/user-redirect/tree or 
  - replace `lab` in your browser url (indicated above in yellow) by `tree` can press `Enter`.
    
---

---

**Tip**

**How to troubleshoot in case or error?**

- Refresh the page, re-login, or clear your browser cache to reset your browser.
- Restart your server using the `Hub Control Panel` under the `File` menu (shown above) or visit the hub home  
  https://divedeep.cs.cityu.edu.hk/hub/home  
  and click `Stop My Server` and then `Start My Server`.  
  ![Restart my server](images/restart.dio.svg)
- Contact your instructor using the contact information on the [course homepage](https://canvas.cityu.edu.hk/courses/42868).

---

## Jupyter Notebook

### How to access course materials?

The recommended way is to follow the links to the lecture/lab notebooks on the course homepage.

![course homepage](images/canvas.dio.svg)

---

**Important**

It is important to open the notebooks using the links on the course homepage as it 
- [git-pull](https://git-scm.com/docs/git-pull) updates/corrections to the notebooks and 
- merge them with the notebooks stored under `cs1302` folder in your home directory but
- without overwriting your changes.

If you open the notebook directly on JupyterHub instead of following the git-pull link, updates will not be pulled, and your notebook may fail to have the updates needed for submission.

If you are interested in how [`nbgitpuller`](https://jupyterhub.github.io/nbgitpuller/index.html) works, the following is an example of a git-pull link:

[$
\begin{alignat*}{2}
&\text{https://divedeep.cs.cityu.edu.hk/hub/user-redirect/git-pull?} &\quad&\leftarrow \text{send request to nbgitpuller app}\\
&\text{repo=https://github.com/ccha23/cs1302\&} &&\leftarrow \text{pull from the repository}\\
&\text{urlpath=lab/tree/cs1302/release/Lab0/Setup.ipynb} &&\leftarrow \text{follow the path to open the file}
\end{alignat*}
$](https://divedeep.cs.cityu.edu.hk/hub/user-redirect/git-pull?repo=https%3A%2F%2Fgithub.com%2Fccha23%2Fcs1302&urlpath=lab%2Ftree%2Fcs1302%2FLab0%2FSetup.ipynb)

---

### How to complete a lab assignment?

Learn how to edit a notebook:

1. Click `Help`$\to$ `Welcome Tour` to learn the JupyterLab interface.  
2. Click `Help`$\to$ `Notebook Tour` to learn about the jupyter notebook.
3. Watch the official [video tutorial](https://youtu.be/A5YyoCKxEOU) to JupyterLab.

**Exercise** 

In learning a new programming language, the first program to write is often the ["Hello, World!"](https://en.wikipedia.org/wiki/%22Hello,_World!%22_program) program, which says Hello to the world.  

Type the program below and run it with `Shift+Enter`. Use tab or 4 spaces to indent.

![Hello world program](images/hello.dio.svg) 

<!-- grade_id: helloworld, points: 0, grade: false, solution: true -->
```python
def say_hello():
    # YOUR CODE HERE
    raise NotImplementedError()
    
say_hello()
```

---

**Attention**

We often ask you to write your answer/code in a special cells with *unique but hidden id number*. 
- Make sure you provide your answers only to the answer cells that says `YOUR CODE HERE` or "YOUR ANSWER HERE". 
- For coding exercises, you should also remove `raise NotImplementedError()`.
- Do not clone or duplicate the answer cells.

---

In order to check your work thoroughly, there will be both *visible* and *hidden* test cases. The following is a visible test you can run to check your answer:

<!-- grade_id: test-helloworld, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your "Hello, World!" program.
import io
import sys

old_stdout, sys.stdout = sys.stdout, io.StringIO()
say_hello()
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
assert printed == "Hello, World!\n"
```

The test returns an assertion error only if your program does not print the correct message.

---

**Tip**

1. You can repeatedly modify your solution and run the test cell until your solution passes the test. There is no mark penalty in failing the test before submission. 
1. The test cell often uses more tools that what you might have learned. You are not required to know how to write the test cell but you should try to understand the error messages from a failed test.
1. To assess your solution thoroughly, we will run new tests hidden from you after you have submitted your notebook. There is no partial credit for a partially correct solution that works for the visible test but fails for the hidden test. Therefore, *you should ensure your solution works in general rather than just the visible tests*.
1. If you open the same notebook multiple times in different browser windows, be careful in making changes in different windows as you may overwrite your own changes.
1. If your notebook fails to run any code, the Kernel might have died. You can restart the kernel with `Kernel`$\to$ `Restart`. If restarting fails, check your code cells to see if there is any code that breaks the kernel.

---

### How to submit a notebook (Important!)

1. Take a screenshot for each answer. 
2. Paste them in a word file and convert it to be a pdf file.
3. Re-name your pdf file to be student_number.pdf (change student number to your number, e.g., 123456.pdf)
4. Submit the pdf file to Canvas-Quiz before the deadline.

**Although this lab assignment is not counted to your final mark, it's strongly recommended that you submit it in the Canvas to avoid any issues.**
**Failing to obey the rules above may make your submission invalid.**

```python

```